# Einsendeaufgabe 1: Numerische Genauigkeit und Gleichungssysteme (100 Punkte)
In dieser Aufgabe sollen Sie ein wenig mehr Erfahrung mit NumPy und numerischen Methoden gewinnen.  
Zur Erinnerung empfehle ich an dieser Stelle, die Definition der [IEEE-Fließkommazahlen](https://de.wikipedia.org/wiki/IEEE_754) zu wiederholen.  

## Addition von Zahlen (20 Punkte) 

Gegeben sei ein Array *array*, dass 100 mal die Zahl $10^{-16}$ enthält und einmal (als ersten Eintrag) die Zahl $1$. 

```python
import numpy as np

array = np.concatenate(([1], np.full(100, 1e-16)))
print(array)
```

Output:
```
[1.e+00 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16
 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16
 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16
 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16
 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16
 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16
 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16
 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16
 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16
 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16
 1.e-16]
```

**Aufgabe:** Addieren Sie alle Werte in `array`:
- mit der Funktion `np.sum()`
- mit einer Schleife, die von **vorne nach hinten** über `array` iteriert
- mit einer Schleife, die von **hinten nach vorne** über `array` iteriert

_Points:_ 6

```python
sum_numpy = array.sum()
print(f"{sum_numpy=  }")

sum_forward = 0
for value in array:
    sum_forward += value
print(f"{sum_forward=}")

sum_reverse = 0
for value in array[::-1]:
    sum_reverse += value    
print(f"{sum_reverse=}")
```

Output:
```
sum_numpy=  np.float64(1.0000000000000084)
sum_forward=np.float64(1.0)
sum_reverse=np.float64(1.00000000000001)
```

<!-- BEGIN QUESTION -->

**Aufgabe:** Erklären Sie die Ergebnisse. Wie werden die Zahlen 1 und 1e-17 im Computer dargestellt? (ausführliche Erläuterung erforderlich!)  

_Points:_ 12

Die Praezision einer Gleitpunktzahl haengt von der Mantisse ab. Diese bildet neben dem Exponenten, der Basis und dem Vorzeichen die Gleitpunktzahl. 

Solange man nicht explizit davon abweicht, wird float64 in Python und numpy verwendet fuer Gleitpunktzahlen.

Eine float32-Mantisse besteht aus 23 Bit und eine float64-Mantisse besteht aus 52 Bit (bleiben bei 64 Bit noch 11 Bit fuer den Exponent und ein Bit fuer das Vorzeichen).

Werden Werte die zur Darstellung mehr als die in der Mantisse vorhandenen Bits benoetigen dargestellt, kommt es zu Rundungsfehlern wie unten erlaeutert. 

```python
# x + y nur darstellbar wenn abs(np.log2(x)-np.log2(y)) < Anzahl Mantissenbits
def mindest_mantissenbits(x,y):
    #berechnet die Anzahl der benoetigten Mantissenbits fuer eine korrekte Berechnung
    #float128 um exaktere Praezision zu ermoeglichen
    return np.float128(abs(np.float128(np.float128(np.log2(np.float128(x)))-np.float128(np.log2(np.float128(y))))))
```

```python
#Beispiel zur Darstellbarkeit float32
#hoechste darstellbare Zahl in 23bit, passt noch in 23bit Mantisse
x = 1
y = np.float128(2**24-2)
x+y, np.float32(x+y), "darstellbar:", mindest_mantissenbits(x,y) < 23 + 1
```

Output:
```
(16777215.0, 16777215.0, 'darstellbar:', True)
```

```python
#Beispiel zur Darstellbarkeit float32
#Mantisse = 1 wenn dargestellte Zahl Potenz von 2, also darstellbar
x = 1
y = np.float128(2**24-1)
x+y, np.float32(x+y), "darstellbar:", mindest_mantissenbits(x,y) < 23 + 1
```

Output:
```
(16777216.0, 16777216.0, 'darstellbar:', True)
```

```python
#Beispiel zur Darstellbarkeit float32
#Dargestellte Zahl muss hier mit Mantisse von Mindestlaenge 24 dargestellt werden, nicht moeglich bei float32
x = 1
y = np.float128(2**24)
x+y, np.float32(x+y), "darstellbar:", mindest_mantissenbits(x,y) < 23 + 1
```

Output:
```
(16777217.0, 16777216.0, 'darstellbar:', False)
```

```python
#Beispiel zur Darstellbarkeit float64
#hoechste darstellbare Zahl in 52bit, passt noch in 52bit Mantisse
x = 1
y = np.float128(2**53 - 2) 
x+y, np.float64(x+y), "darstellbar:", mindest_mantissenbits(x,y) < 52 + 1
```

Output:
```
(9007199254740991.0, 9007199254740991.0, 'darstellbar:', True)
```

```python
#Beispiel zur Darstellbarkeit float64
#Dargestellte Zahl muss hier mit Mantisse von Mindestlaenge 53 dargestellt werden, nicht moeglich bei float64
x = 1
y = np.float128(2**53)
x+y, np.float64(x+y), "darstellbar:", mindest_mantissenbits(x,y) < 52 + 1
```

Output:
```
(9007199254740993.0, 9007199254740992.0, 'darstellbar:', False)
```

```python
#Beispiel 1 + 1e-16, vorkommend bei Summierung in Schlefe von vorne
x = 1
y = np.float128(1e-16)
x+y, np.float64(x+y), "darstellbar:", mindest_mantissenbits(x,y) < 52 + 1
```

Output:
```
(1.0000000000000001, 1.0, 'darstellbar:', False)
```

```python
#Beispiel 1 + 1e-16 * 100, vorkommend bei Summierung in Schlefe von hinten
x = 1
y = np.float128(1e-16 * 100)
x+y, np.float64(x+y), "darstellbar:", mindest_mantissenbits(x,y) < 52 + 1
```

Output:
```
(1.00000000000001, 1.00000000000001, 'darstellbar:', True)
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Aufgabe:** Welches Verfahren verwendet `np.sum()` für die Addition?

_Points:_ 2

In den Numpy docs steht theoretisch: "often numpy will use a  numerically better approach (partial pairwise summation) leading to improved precision in many use-cases".

Hierbei werden die Werte paarweise addiert wodurch eher gleich hohe Werte miteinander addiert werden.

Im Gegensatz dazu steht das Summieren in der Schleife wo stets der Hohe Summenwert mit den geringen Arraywerten addiert wird.

Ein Funktionsentwurf welcher paarweise addiert ist im Folgenden angegeben.

```python
def paarweise_sum(arr):
    """array in gleich lange Teile teilen und diese beiden summieren""" 
    arr=arr.copy()
    last = None    
    while arr.size > 1:
        if arr.size % 2 == 1:            
            arr[0] += arr[-1]
            arr = arr[:-1]        
        arr = arr[::2] + arr[1::2] 
    return arr[0]
paarweise_sum(array)
```

Output:
```
1.0000000000000098
```

Tatsaechlich verwenden numpy jedoch np.add.reduce(array).

Dieses addiert die Werte des Array laut docs von Vorne.

```python
np.add.reduce(array)
```

Output:
```
1.0000000000000084
```

```python
help(np.add.reduce)
```

Output:
```
Help on built-in function reduce:

reduce(...) method of numpy.ufunc instance
    reduce(array, axis=0, dtype=None, out=None, keepdims=False, initial=<no value>, where=True)

    Reduces `array`'s dimension by one, by applying ufunc along one axis.

    Let :math:`array.shape = (N_0, ..., N_i, ..., N_{M-1})`.  Then
    :math:`ufunc.reduce(array, axis=i)[k_0, ..,k_{i-1}, k_{i+1}, .., k_{M-1}]` =
    the result of iterating `j` over :math:`range(N_i)`, cumulatively applying
    ufunc to each :math:`array[k_0, ..,k_{i-1}, j, k_{i+1}, .., k_{M-1}]`.
    For a one-dimensional array, reduce produces results equivalent to:
    ::

     r = op.identity # op = ufunc
     for i in range(len(A)):
       r = op(r, A[i])
     return r

    For example, add.reduce() is equivalent to sum().

    Parameters
    ----------
    array : array_like
        The array to act on.
    axis : None or int or tuple of ints, optional
        Axis or axes along which a reduction is performed.
        The default (
```

<!-- END QUESTION -->

## Grenzwerte von Funktionen (20 Punkte)

Geben sei die Funktion $f(x)=\frac{e^x-1}{x}$

Es gilt: $lim_{x \to 0}f(x) = 1$

**Aufgabe:** Schreiben Sie eine Python Funktion `f(x)`, mit deren Hilfe Sie $f(x)$ für $x= 10^{-1}, 10^{-2}, \cdots , 10^{-15}$ mit NumPy berechnen und geben Sie das Ergebnis aus.

<!-- grade_id: cell-d5ed715a9745e7d6, points: 0, grade: false, solution: true -->
```python
import numpy as np


def f(x):
    return (np.e**x - 1) / x if x != 0 else "undefiniert"


input_numbers = [10**-x for x in range(1,16)]
result_limits = [f(x) for x in input_numbers]
print(f"{result_limits=}")
```

Output:
```
result_limits=[1.0517091807564771, 1.005016708416795, 1.0005001667083846, 1.000050001667141, 1.000005000006965, 1.0000004999621837, 1.0000000494336803, 0.999999993922529, 1.000000082740371, 1.000000082740371, 1.000000082740371, 1.000088900582341, 0.9992007221626409, 0.9992007221626409, 1.1102230246251565]
```

**Aufgabe:** Berechnen Sie nun die Funktion für  $x= 10^{-16}, 10^{-17}, \cdots , 10^{-20}$

_Points:_ 8

<!-- grade_id: cell-86fd557d439b5a41, points: 0, grade: false, solution: true -->
```python
input_numbers_2 = [10**-x for x in range(16,21)]
result_limits_2 = [f(x) for x in input_numbers_2]
print(f"{result_limits_2=}")
```

Output:
```
result_limits_2=[0.0, 0.0, 0.0, 0.0, 0.0]
```

<!-- BEGIN QUESTION -->

**Aufgabe:** Was fällt auf? Welche Erklärung haben Sie für das Ergebnis? (Ausführliche Erklärung erforderlich) 

_Points:_ 12

Die Resultate lagen faelschlicherweise bei 0 weil der Zahler $e^x-1$ ab hier bei 0 liegt.

Da die Mantisse bei float64 nur 52 Bit lang ist, koennen Veraenderungen von unter $\frac {1}{2^{52}}$ nicht erfasst werden. 

Es ergiebt sich also eine minimale Schrittweite von $2^{-52}$.

```python
#der Zaehler e**x - 1 ist hier Null
[(np.e**x - 1, x) for x in input_numbers_2]
```

Output:
```
[(0.0, 1e-16), (0.0, 1e-17), (0.0, 1e-18), (0.0, 1e-19), (0.0, 1e-20)]
```

```python
#bei np.e**1e-15 - 1 ist der Zaehler noch ueber der minimalen Schrittweite von 2**-52
np.e**1e-15 - 1 >= 2**-52, 0 + (np.e**1e-15 - 1)
```

Output:
```
(True, 1.1102230246251565e-15)
```

```python
#bei np.e**1e-16 - 1 ist der Zaehler unterhalb der minimalen Schrittweite von 2**-52
np.e**1e-16 - 1 >= 2**-52, 0 + (np.e**1e-16 - 1)
```

Output:
```
(False, 0.0)
```

<!-- END QUESTION -->

## Gauß-Verfahren und LU-Zerlegung (60 Punkte) 
Obwohl man, wie wir sehen werden, lineare Gleichungssysteme einfach mit SciPy lösen kann, wollen wir in dieser Aufgabe  einmal selbst eine sogenannte LU-Zerlegung bzw. das Gauß-Verfahren implementieren. Außerdem werden wir diesen Code auch später noch verwenden, wenn wir uns mit Performance-Optimierung und Parallelisierung beschäftigen.

Das Gauß-Verfahren ist ein wichtiges Verfahren zum Lösen von linearen Gleichungssystemen, wir gehen hier schrittweise mit der Implementierung vor.

**Aufgabe:** Implementieren Sie die Funktion `b_solve` zur Lösung einer Dreiecksgleichung der Form:

$\left[ {\begin{array}{cccc} 
a_{11} & a_{12} & \cdots & a_{1n}\\ 
0 & a_{22} & \cdots & a_{2n}\\ 
\vdots & \vdots & \ddots & \vdots \\ 
0 & 0 & \cdots & a_{nn}\\ 
\end{array} } \right] \left[{\begin{array}{c}  
x_1\\ 
x_2\\ 
\vdots\\ 
x_n\\ 
\end{array}}\right] =  
\left[{\begin{array}{c}  
b_1\\ 
b_2\\ 
\vdots\\ 
b_n\\ 
\end{array}}\right]$ 

Hier beginnt man mit der Löung der letzten Zeile, denn es gilt:

$x_n = b_n/a_{nn} $
und dann

$x_{n-1} =  (b_{n-1}-a_{(n-1),n}*x_{n})/a_{(n-1),(n-1)}$

$x_{n-2} =  (b_{n-2}-(a_{(n-2),n}*x_{n}+a_{(n-2),(n-1)}*x_{n-1})/a_{(n-2),(n-2)}$


*Hinweis*: Achten Sie darauf, dass ihr Code nicht nur korrekt, sondern auch effizient ist. 

_Points:_ 12

<!-- grade_id: cell-045bae77f313bff4, points: 0, grade: false, solution: true -->
```python
import numpy as np

def b_solve(A, b, upper_triangular=True):
    """Solve a system of linear equations A * x = b where A is an upper triangular matrix."""
    #kann auch nicht-triangulare Matrizen loesen 
    
    #Gauss-Jordan-Algorithmus    
    Ab=np.concatenate((A,b),1)    
    for n in range(min(A.shape)-1,-1,-1):          
        if Ab[n][n]==0:
            continue
        Ab[n]=Ab[n]/Ab[n][n]  
        
        if upper_triangular:
            #Dreiecksmatrix ohne extra Schleife, 11.5 microsec
            Ab[:n] = Ab[:n] - np.outer(Ab[:n,n], Ab[n])
            
            #Dreiecksmatrix, 10 microsec
            #for m in range(n):                    
             #   Ab[m] = Ab[m] - Ab[m][n] / Ab[n][n] * Ab[n]
        else:
            #Allgemeinfall ohne extra Schleife, 20.5 microsec
            zeilen = np.arange(Ab.shape[0]) != n
            Ab[zeilen] = Ab[zeilen] - np.outer(Ab[zeilen, n], Ab[n])

            #Allgemeinfall, 14.5 microsec
            #for m in range(Ab.shape[0]):        
             #   if m != n:            
              #      Ab[m] = Ab[m] - Ab[m][n] / Ab[n][n] * Ab[n]        
        
    return Ab[:,-1:] 
```

```python
#%%timeit
b_solve(np.array([[4,1,1],
                [0,4,3],
                [0,0,4]]),

    np.array([[4.        ],
           [4.87298335],
           [1.12        ]]))
```

Output:
```
array([[0.67793854],
       [1.00824584],
       [0.28      ]])
```

**Aufgabe:** Implementieren Sie die LU-Zerlegung einer Matrix mit Pivotisierung.  
Die Funktion `lu_decomposition` soll die Matrizen L und U berechnen.  
Zu Beginn wird der Funktion eine Matrix A übergeben.  


Implementieren Sie die Funktion "in-place", d.h. die Ergebnisse für L und U werden direkt in A geschrieben.  
Dabei wird die ursprüngliche Matrix überschrieben.  
Die Funktion gibt den Pivot-Vektor zurück, der die folgende Form haben sollte:

$ 
P=\left[{\begin{array}{ccc} 
0 & 1& 0\\ 
0 & 0& 1\\ 
1 & 0& 0\\ 
\end{array} } \right] \Rightarrow \left[{\begin{array}{c}  
1\\ 
2\\ 
0\\ 
\end{array}}\right] 
$ 

D.h. es wird nur der Index des Wertes in jeder Zeile gespeichert, der nicht Null ist.

Wenn man dies macht, kann man später für die Lösung einfach den Vektor `bc = b[P]` verwenden, um auch $b$ richtig zu pivotisieren.  

Hinweise: 
- Die Pivot-Zeile lässt sich am einfachsten mit `np.argmax` finden. Die Zeile mit dem höchsten Betrag können Sie mit `np.abs` finden.

- Tauschen zweier Elemente in einem NumPy Array: `v[[a, b]] = v[[b, a]]`

- Versuchen Sie, die interne Vektorisierung von NumPy im Eliminationsschritt zu nutzen. Verwenden Sie dazu die Slicing-Syntax, um die innerste Schleife zu ersetzen.

_Points:_ 15

<!-- grade_id: cell-b482047b4df4ac5e, points: 0, grade: false, solution: true -->
```python
def lu_decomposition(A):
    """L und U Matrix werden innerhalb von A Matrix berechnet und der Pivot-Vektor wird ausgegeben"""
    n = A.shape[0]
    P = np.arange(n)

    for k in range(n):
        # Pivotisierung: Index der Zeile mit größtem Betrag in Spalte k finden
        pivot_index = np.argmax(np.abs(A[k:n, k])) + k
        if A[pivot_index, k] == 0:
            raise ValueError("Matrix ist singulaer.")

        # Zeilen tauschen in A
        if pivot_index != k:
            A[[k, pivot_index], :] = A[[pivot_index, k], :]
            P[[k, pivot_index]] = P[[pivot_index, k]]

        # Berechnung L unterhalb der Diagonale
        A[k+1:n, k] /= A[k, k]

        # Berechnung U obere Dreiecksmatrix inklusive Diagonale
        A[k+1:n, k+1:n] -= np.outer(A[k+1:n, k], A[k, k+1:n])

    return P #Pivot-Vektor
```

<!-- BEGIN QUESTION -->

**Aufgabe:** Schreiben Sie eine Funktion `custom_solve`, welche mit Hilfe von `lu_decomposition` ein lineares Gleichungssystem der Form $A*x=b$ löst.  

Hinweis: Versuchen Sie, die interne Vektorisierung von NumPy zu nutzen. Verwenden Sie dazu die Slicing-Syntax, um die innerste Schleife zu ersetzen.

_Points:_ 12

<!-- grade_id: cell-e08e49ee4f4639ff, points: 0, grade: false, solution: true -->
```python
def custom_solve(A, b):
    """Solve a system of linear equations A * x = b and return result vector x."""
    
    #Pivot-Vektor und A (welches LU-Zerlegung enthaelt) werden berechnet
    P = lu_decomposition(A)
    
    #L und U werden aus A isoliert
    n = A.shape[0]
    L = np.tril(A, k=-1) + np.eye(n)  # untere Dreiecksmatrix mit Einsen auf der Diagonale
    U = np.triu(A)                   # obere Dreiecksmatrix inklusive Diagonale
    
    #Lc = b[P] wird geloest fuer c
    c = b_solve(L, b[P], upper_triangular=False)
    
    #Ux = c wird geloest fuer x
    x = b_solve(U, c)
    
    return x
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Aufgabe:** Schreiben Sie eine Testfunktion, die überprüft, ob die Lösung ihrer Implementierung des Gauß-Verfahrens richtig ist.  
Testen Sie Gleichungssysteme der Größe $16 \times 16$, $32\times 32$ und $64 \times 64$.

Sie können zum Vergleich die SciPy Funktion `linalg.solve` verwenden:

_Points:_ 6

```python
# Beispiel linalg
import scipy.linalg as linalg

n = 128
rng = np.random.default_rng(seed=123)
A = rng.random((n, n))
b = rng.random(n)
c = linalg.solve(A, b)
```

<!-- grade_id: cell-f6a6bc3ba063c4a4, points: 0, grade: false, solution: true -->
```python
import scipy.linalg as linalg

def test_custom_solve(n):    
    """vergleicht Ergebnisse von custom_solve mit linalg.solve"""
    
    rng = np.random.default_rng(seed=123)
    A = rng.random((n, n))
    #Spaltenvektor b ist 2-dimensional
    b = rng.random((n,1))
    linalg_x = linalg.solve(A, b)    
    custom_x = custom_solve(A, b)
    
    #sicherstellen, dass Ergebnisse gleich sind, float-Unpraezision ignorieren
    assert np.allclose(linalg_x, custom_x, atol=1e-6), f"Ergebnisse weichen zu stark ab bei Dimension {n}"
    
    print("Test bestanden fuer Dimension", n)
    
test_custom_solve(16)
test_custom_solve(32)
test_custom_solve(64)
```

Output:
```
Test bestanden fuer Dimension 16
Test bestanden fuer Dimension 32
Test bestanden fuer Dimension 64
```

<!-- END QUESTION -->

**Aufgabe:** Wir wollen nun weiter die *fertigen* SciPy Funktionen betrachten. Mit den Funktionen `lu` und `lu_factor` aus dem Packet `scipy.linalg `  können Sie eine LU-Dekomposition durchführen [[1]](https://docs.scipy.org/doc/scipy/reference/generated/scipy.linalg.lu_factor.html#scipy.linalg.lu_factor) und [[2]](https://docs.scipy.org/doc/scipy/reference/generated/scipy.linalg.lu.html).  
Beide Funktionen haben unter anderem den Parameter `overwrite_a`, der `False` oder `True` sein kann. 

Evaluieren Sie die Performance (Laufzeit) der LU-Zerlegung einer 4096x4096 Matrix, wenn Sie jeweils die Funktionen `lu_factor` und `lu` mit `overwrite_a = True` oder `overwrite_a = False` verwenden.

In der Liste `minimums` sollen die minimalen Laufzeiten in folgender Reihenfolge stehen:
1. `lu_factor` mit `overwrite_a = False`
2. `lu_factor` mit `overwrite_a = True`
3. `lu` mit `overwrite_a = False`
4. `lu` mit `overwrite_a = True`

Hinweis: Da die Matrix A bei der Verwendung von `overwrite_a = True` verändert wird, kann es zu einer Vereinfachung der Berechnung kommen, wenn die Funktion mehrfach nacheinander ausgeführt wird. Aus diesem Grund muss die `timeit` Funktion [repeat](https://docs.python.org/3/library/timeit.html#timeit.repeat) mit `repeat = 2` und `number = 1`  verwendet werden. Die Matrix wird im `setup` Parameter erstellt (aus Zufallswerten mit dem vorgegebenen "random number generator").

_Points:_ 6

<!-- grade_id: cell-5978db060fba8a0d, points: 0, grade: false, solution: true -->
```python
from scipy.linalg import lu_factor, lu
from timeit import repeat

n = 4096
number = 1
repetitions = 2
# store all four minimal runtimes in this list
rng = np.random.default_rng(seed=123)
repeat_lambda = lambda x: repeat(x, setup="A = rng.random((n, n))", repeat=repetitions, number=number, globals=globals())
minimums = [min(repeat_lambda(x)) for x in ["lu_factor(A, overwrite_a=False)", "lu_factor(A, overwrite_a=True)", "lu(A, overwrite_a=False)", "lu(A, overwrite_a=True)"]]
minimums
```

Output:
```
[9.591415676986799, 9.58492271893192, 9.971174363978207, 9.970445337006822]
```

**Aufgabe:** Untersuchen Sie den Speicherverbrauch der 4 Varianten mit memray.

Um den Einfluss der anderen Berechnungen zu minimieren, empfehlen wir hier, die Berechnungen in eine neue Datei zu schreiben und mit `!` in einer Code-Zelle memray's Command-Line Interface zu verwenden. 

Am besten tracen Sie nur die LU Funktion mit einem Context Manager!
Verwenden Sie jeweils folgende Dateinamen als `file_name` Parameter in `memray.Tracker`:

- lu_inplace.bin
- lu_notinplace.bin
- lu_factor_inplace.bin
- lu_factor_notinplace.bin

Beantworten Sie bitte die folgenden Fragen:
* Welche Version verwendet am meisten Speicher? 
* Bei welcher Version ist der Speicherverbrauch am konstantesten? 

_Points:_ 9

```python
#.py files werden erstellt fuer die jeweiligen Funktionen
for file_name in ["lu_inplace.bin", "lu_notinplace.bin", "lu_factor_inplace.bin", "lu_factor_notinplace.bin"]:
    with open(file_name.replace(".bin", ".py"), "w") as file:
        #ich muss den Pfad des Environment im dem mein memray liegt hinzufuegen
        file.write(
"""import sys
base_path = "/home/afu/Programmieren/python_sachen/"
sys.path.append(base_path + "myenv/lib/python3.12/site-packages/")
del sys
""")       
        
        file.write(
f"""import numpy as np
from scipy.linalg import lu_factor, lu
import memray

n = 4096
rng = np.random.default_rng(seed=123)
A = rng.random((n, n))

with memray.Tracker("{file_name}", native_traces=True):
    lu{"_factor" if "_factor" in file_name else ""}(A, overwrite_a={"_inplace" in file_name})
""")
    
```

```python
#ich muss den Pfad des Environment im dem mein memray liegt hinzufuegen
import sys
base_path = "/home/afu/Programmieren/python_sachen/"
sys.path.append(base_path + "myenv/lib/python3.12/site-packages/")
del sys

import memray, os
```

```python
#alle .py files werden ausgefuehrt, bei mir mit python3
for file_name in ["lu_inplace.bin", "lu_notinplace.bin", "lu_factor_inplace.bin", "lu_factor_notinplace.bin"]:
    os.system(f"python3 {file_name.replace(".bin", ".py")}")
```

Output:
```
Memray WARNING: Correcting symbol for malloc from 0x420620 to 0x7c2201ead650
Memray WARNING: Correcting symbol for free from 0x420ab0 to 0x7c2201eadd30
Memray WARNING: Correcting symbol for malloc from 0x420620 to 0x7756e32ad650
Memray WARNING: Correcting symbol for free from 0x420ab0 to 0x7756e32add30
Memray WARNING: Correcting symbol for malloc from 0x420620 to 0x7a401a2ad650
Memray WARNING: Correcting symbol for free from 0x420ab0 to 0x7a401a2add30
Memray WARNING: Correcting symbol for malloc from 0x420620 to 0x76cbab2ad650
Memray WARNING: Correcting symbol for free from 0x420ab0 to 0x76cbab2add30
```

```python
#Speichernutzung wird aus allen .bin files gelesen
for file_name in ["lu_inplace.bin", "lu_notinplace.bin", "lu_factor_inplace.bin", "lu_factor_notinplace.bin"]:
    reader = memray.FileReader(file_name)
    mem_use = 0

    for allocation in reader.get_allocation_records():
        mem_use += allocation.size

    print(f"Speichernutzung von {file_name[:-4]} : {round(mem_use / 2**20)}MB") 
```

Output:
```
Speichernutzung von lu_inplace : 400MB
Speichernutzung von lu_notinplace : 528MB
Speichernutzung von lu_factor_inplace : 144MB
Speichernutzung von lu_factor_notinplace : 144MB
```

Die Funktion lu_factor verwendet konstant nur 144MB an Speicher. 

Funktion lu verwendet deutlich mehr Speicher mit 400MB "inplace" und sogar 528MB "notinplace"

```python
#dateien wieder loeschen
#!rm -rf lu_* memray-*
```

```python

```

Tests for addition_ieee (points: 12)


Tests for addition_npsum (points: 2)


Tests for addition_variants (points: 6)


Tests for limits_explanation (points: 12)


Tests for limits_fx (points: 8)
```
>>> print(result_limits)
[1.05170918 1.00501671 1.00050017 1.00005    1.000005   1.0000005
 1.00000005 0.99999999 1.00000008 1.00000008 1.00000008 1.0000889
 0.99920072 0.99920072 1.11022302]
```
```
>>> print(result_limits_2)
[0. 0. 0. 0. 0.]
```


Tests for linearsystem_bsolve (points: 12)


Tests for linearsystem_customsolve (points: 12)


Tests for linearsystem_customsolve_test (points: 6)


Tests for linearsystem_ludecomp (points: 15)


Tests for linearsystem_scipy_memory (points: 9)


Tests for linearsystem_scipy_time (points: 6)
